```python
# Initialize Otter
import otter
grader = otter.Notebook("intro2Class.ipynb")
```

## Lecture Section

Welcome to XXXXXX!

In this lecture, we will cover:
* Lecture-assignments
* The auto-grader
* Submitting assignments

This course is an introduction to Python from a statistical perspective. All assignments will be accessed through GitHub and uploaded to Gradescope for grading. Gradescope connects to D2L, where you will find lectures, supplemental learning content, and most importantly: the syllabus. Read it carefully if you haven't already. Additionally, you should have read the "Getting Started" webpage on D2L, which sets you up with the necessary programs and knowledge to begin class.

If you are reading this, it's safe to say that you've successfully grabbed your assignment from GitHub. If you haven't, revisit the "Getting Started" page. You will always grab an assignment this way, and they will always be linked to their respective assignment folders on D2L, too. All assignments will be completed in a Jupyter notebook. Using GitHub allows you to use GitHub Codespaces - but if you prefer an IDE or already have one installed (I like PyCharm & VSCode) - feel free to use those.

Once you've accessed an assignment, you will notice that they are split into a 'Lecture' and an 'Assignment' section. The 'Lecture' section provides a tutorial-esque walk-through of the relevant lecture. The 'Assignment' section is your homework.

Since this course is asynchronous, there are no due-dates. However, due-dates have been provided in the syllabus for those of us that need structure. The course can be completed in a typical 14/15- week semester, or you can take it at an accelerated 7-week pace.

A single assignment should take you 30 minutes in the 'Lecture' section, and about an hour in the 'Assignment' section. If you are aiming for a full-semester, you should aim for two assignments a week. An accelerated pace is double that, or four assignments a week.

Here is the fun part: this class uses an autograder to grade your work. This allows you a lot of flexibility in submitting assignments. The autograder is coded into each assignment, including this one. Don't worry - if there's ever a problem, you can ask your instructor for help. For now, though, I am going to show you how to use it.


The autograder that this course utilizes is called Otter. It is a modular open-source autograder developed by the Data Science Education Program at UC Berkeley. Since D2L doesn't connect to most external learning tools, we have a unique pipeline for assignments.

1. Use our GitHub Classroom to access your assignments. You can choose to code within GitHub Codespaces, or you can clone your repository and use an IDE of your choice. If you are reading this, this step should be complete!
2. Work on your assignment. Frequently push your assignment to your assignment repository. Make sure your assignment repository is up-to-date.
3. When you are ready to submit, head over to Gradescope. Gradescope connects to D2L and allows automatic autograding using Otter. Choose your assignment, the choose Github as the submission method.
4. You can submit multiple times before the deadline.
5. Gradescope will show you which tests you passed, and which you failed. This class does not use hidden tests - the public tests you see are all the will be couunted toward your grade.


Other specifics aren't important, but I've linked the documentation below, if you are interested:

https://otter-grader.readthedocs.io/en/latest/index.html

Most importantly, you can run the same tests locally on your own machine. You will notice an import statement, and a grader variable. These lines allow you to use the autograding locally. The tests are saved in the metadata of your assignment file. To check a particular question, you run:

grader.check("x")

The "x" will be automatically provided, and is the underlying label of the question. If you accidentally delete the cell, it is likely the "x" will be "qk" where "k" is the question number you are working on.

If you want to check your entire document, you can run:

grader.check_all()

Make sure you do so in a new cell, and not in a cell that contains your solution.

We will test these functions in a few moments.


If a question cannot be graded by the autograder, it will be labelled (manual). The system will not grade these questions, but your instructor will.

Finally, I want to make sure you are aware of the course Piazza link. If there is a problem with the autograder, GitHub, Gradescope, or any other piece of our pipeline, please post there first. To join, click the link below:

## Assignment Section


**Question 1.** You do not need to do anything for this question, except to test that the autograder is working appropriately. This question should pass it's test-case. To run it, click the green arrow (or its equivalent) next to the code block. Make sure to run all code-blocks before it, first.


```python
def sieve(n):
    """
    Generate a set of prime numbers less than or equal to a positive integer.
    """
    is_prime = [True for _ in range(n + 1)]
    p = 2
    while p ** 2 <= n:
        if is_prime[p]:
            for i in range(p ** 2, n + 1, p):
                is_prime[i] = False
        p += 1

    is_prime[0]= False
    is_prime[1]= False

    return set(i for i in range(n + 1) if is_prime[i])
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2.** Once again, make sure the autograder passes this function. Once done, run the 'grader.check_all()' function and make sure it works, too. The assignment can be pushed to GitHub and submitted to Gradescope if all works appropriately.

```python
def hailstone(n):
    """
    Generate the hailstone sequence of a positive integer.
    """
    if n == 1:
        return [n]
    elif n % 2 == 0:
        return [n] + hailstone(n // 2)
    else:
        return [n] + hailstone(3 * n + 1)

```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!
```

Tests for q1 (points: None)
```
>>> assert sieve(1) == set()
>>> assert sieve(2) == {2}
>>> assert sieve(3) == {2, 3}
```
```
>>> assert sieve(20) == {2, 3, 5, 7, 11, 13, 17, 19}
>>> assert sieve(100) == {2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71, 73, 79, 83, 89, 97}
```


Tests for q2 (points: None)
```
>>> assert hailstone(1) == [1]
>>> assert hailstone(2) == [2, 1]
>>> assert hailstone(4) == [4, 2, 1]
```
```
>>> assert hailstone(20) == [20, 10, 5, 16, 8, 4, 2, 1]
>>> assert hailstone(9) == [9, 28, 14, 7, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
```
